Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "Batsal Ghimire"
COLLABORATORS = ""
```

---

# CS110 Pre-class Work 2.2

## Question 1 (Exercise 3.1-3 of Cormen, et al. )
Explain why the statement, "The running time of algorithm A is at least $O(n^2)$," is meaningless.


<!-- grade_id: cell-437d61c1420d4f99, points: 0, grade: true, solution: true -->
Big O describes the worst case scenario of the algorithm, and is the upper bound for the runtime of the algorithm. And from the statement we can see that it is describing the lower bound (i.e. at least) which is not the proper use of the  notation. Also, there is no information given about the worst-case scenario which is essential while analyzing an algorithm. Rather than using the phrase "at least $O(n^2)$", there is a separate notation made for such use i.e. big-$\Omega$, which can be used.

## Question 2 (Exercise 3.1-4 of Cormen, et al. )

Is $2^{n+1}=O(2^n)$? Is $2^{2n}=O(2^n)$?

<!-- grade_id: cell-6a0c4ee19dadfddf, points: 0, grade: true, solution: true -->
Yes, $2^{n+1} = O(2^n)$. This is because writing $2^{n+1}$ is the exactly same as writing $2^n.2^1 = 2.2^n$. And we know that we are not concerned with the constant factor since we're analyzing order of growth. So, we can omit 2 from the big-O notation.

However, if we break the given function $2^{2n}$, we get, $2^n.2^n$. For the given inequality to be true, there needs to be a constant, c where: $2^{2n} = c.2^n$. As 'n' increases exponentially, we can see that there exists no such value rendering the equation false.

## Question 3.
Write a function in Python that solves the maximum-subarray problem using a brute-force approach. Your Python function must:
* Take as Input an array/list  of numbers
* Produce the following Output: 
    * the start and end indices of the subarray containing the maximum sum.
    * value of the maximum subarray (float)


<!-- grade_id: cell-527e6532d6992fd0, points: 0, grade: false, solution: true -->
```python
import math
def bruteforce_max_subarray(A):
    """Implements brute-force maximum subarray finding.
    
    Inputs:
    - A: a NON-EMPTY list of floats
    
    Outputs: A tuple of
    - the start index of the max subarray
    - the end index of the max subarray
    - the value of the maximum subarray
    """
    if len(A)>1:
        n = len(A) # Sets 'n' to be the length of the array A
        max_sum = -math.inf # Since we have negative numbers, we need to set the highest_sum as -infinity to ensure that it is the lowest number.
        for i in range(n): # Goes through all the element of the list.
            for j in range(i+1,n): # We assuume that the sum of the same element with itself does not exist.
                if max_sum < sum(A[i:j+1]): #Compares the max_sum to the sum of the elements in the given range. The j+1 is used because in python the end range is not included.
                    max_sum = float(sum(A[i:j+1])) #If the new sum is greater than previous max_sum, the max_sum is replaced with this new higher value.
                    index_start = i # Sets the initial index to
                    index_end = j # Sets the last index to index_end
        return (index_start, index_end, max_sum) # Returns the start index, end index and value of maximum subarray.
        raise NotImplementedError()
    else:
        print ("There is only one element.")
```

<!-- grade_id: cell-a28a56466c9537ff, points: 1, grade: true, solution: false -->
```python
assert(bruteforce_max_subarray([-2,1,-1,2,-5]) == (1, 3, 2))
assert(bruteforce_max_subarray([-2, -5, 6, -2, -3, 1, 5, -6]) == (2, 6, 7))
```

## Question 4. 
Test your Python maximum-subarray function using the following input list (from Figure 4.3 of Cormen et al.):  
`A = [13, -3, -25, 20, -3, -16, -23, 18, 20, -7, 12, -5, -22, 15, -4, 7] `

If your Python implementation is correct, your code must return: 
* 43 - which is the answer to the maximum subarray problem, and 
* <7, 10> -the start and the end indices of the max subarray. 



<!-- grade_id: cell-e4a632ce0edc1697, points: 0, grade: true, solution: true -->
```python
# YOUR CODE HERE
A = [13, -3, -25, 20, -3, -16, -23, 18, 20, -7, 12, -5, -22, 15, -4, 7]
test = bruteforce_max_subarray(A)
index = test[0:2]
answer = test[2]
print ("The answer is:",answer)
print ("The start and end index is:", index)
assert(bruteforce_max_subarray([13, -3, -25, 20, -3, -16, -23, 18, 20, -7, 12, -5, -22, 15, -4, 7]) == (7,10,43))
# We don't get any assertion error, so the answer is correct.
```

Output:
```
The answer is: 43.0
The start and end index is: (7, 10)
```

## Question 5. Asymptotic notation. 
Complete the following table using the asymptotic notation that best describes the problem. For example, if both $O(n^3)$ and $O(n)$ are possible for an algorithm, the answer is $O(n)$ because the function $f(n) = O(n)$ provides a tighter and more accurate fit; if both $O(n)$ and $\Theta(n)$ are possible, the correct answer is $\Theta(n)$ because $\Theta(n)$ provides both information about the upper and lower bound, thus it is more accurate than $O(n)$.

You should copy the following table and paste and edit it in the cell below. 

Algorithm | Big Oh ($O$) | Big Theta ($\Theta$)
--- | --- | ---
Insertion sort |  |
Selection sort |  |
Bubble sort |  | 
Finding maximum subarray |  |

<!-- grade_id: cell-247c51df276b622e, points: 0, grade: true, solution: true -->
Algorithm | Big Oh ($O$) | Big Theta ($\Theta$)
--- | --- | ---
Insertion sort |$O(n^2)$  |$\Theta (n^2)$
Selection sort |$O(n^2)$  |$\Theta (n^2)$
Bubble sort |$O(n^2)$  |$\Theta (n^2)$
Finding maximum subarray |$O(n^2)$  |$\Theta (n^2)$

## [Optional] Question 6. 
How can you change this code to make it find the minimum-subarray?

<!-- grade_id: cell-f4553f71858d1bc4, points: 0, grade: false, solution: true -->
```python
import math
def bruteforce_min_subarray(A):
    """Implements brute-force maximum subarray finding.
    
    Inputs:
    - A: a NON-EMPTY list of floats
    
    Outputs: A tuple of
    - the start index of the max subarray
    - the end index of the max subarray
    - the value of the maximum subarray
    """
    # YOUR CODE HERE
    n = len(A)
    min_sum = math.inf
    for i in range(n): 
        for j in range(i,n):
            if min_sum > sum(A[i:j+1]): 
                min_sum = sum(A[i:j+1])
                start = i
                end = j
    return (start, end, min_sum)
    raise NotImplementedError()
```

<!-- grade_id: cell-4263c7494a5f09d3, points: 0, grade: true, solution: false -->
```python
assert(bruteforce_min_subarray([1]*10)[0] == bruteforce_min_subarray([1]*10)[1])
assert(bruteforce_min_subarray([1]*10)[2] == 1)
```