Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "Batsal Ghimire"
COLLABORATORS = ""
```

---

# CS110 Pre-class Work 7.1

## Part A. Median Heap (watch this [video explanation](https://www.youtube.com/watch?v=756_8C9YBZQ&list=PLF_a-qBXTGFektoI6JUOTRL36JlvD04BR&index=5&t=0s) or read this [description](https://stackoverflow.com/a/15319593/7946759))

Throughout this pre-class work, please use the following definition of median: the median of a list of numbers is the one in the middle of the list when the list is ordered. When such the middle element can’t be determined (i.e., in a list of even length), the average of the two middle elements is the median. For example, 5 is the median of [-1,2,4,5,8,10,12], and (5+7)/2=6 is the median of [1,2,3,5,7,8,10,11].

Using the idea from Lesson 3.2, we can use a pair of heaps to create a data structure which allows fast access to the median. Use the heapq module in python to create both a max-heap and a min-heap. Note that by default, the heapq module in python only creates min-heaps, but if we multiply elements by -1 when we store them, then we can also create max-heaps.


## Question 1.
Write a function `add_to_median_heap(minh, maxh, elem)`. It must accept a min heap, a max heap, and an element to add.


<!-- grade_id: cell-50cb08de70712ee7, points: 0, grade: false, solution: true -->
```python
import heapq
def add_to_median_heap(minh, maxh, elem):
    if abs(len(maxh) - len(minh)) <= 1: #Ensures that the difference between the length of the two lists are not greater than 1.
        if (len(maxh) == 0) or (len(minh) > len(maxh)): #If the maxh is empty or it is smaller than minh.
            maxh.append(elem) #In that case, add the element to maxh list.
        elif(len(minh) == 0) or (len(maxh) > len(minh)): #If the minh is empty or minh is smaller than maxh.
            minh.append(elem) #Add the element to the minh list.
        else: #If both the list size is equal and not empty.
            if (elem >= maxh[0]) and elem >= minh[0]: #If the element is greater than or equal to the prioritized element from max and min heap.
                maxh.append(heapq.heappop(minh)) #Pop the value from minh and add it to maxh
                minh.append(elem) #Add the new element in the min heap.
            elif (elem <=  minh[0] and elem <= maxh[0]):#If the element is smaller than or equal to the prioritized element from max and min heap.
                minh.append(heapq.heappop(maxh))#Pop the value from the maxh and add it to minh
                maxh.append(elem) #Add the new element in the maxheap.
            elif (elem >=  maxh[0]): #If the new element is greater than or equal to maxh.
                minh.append(elem) #Add the element to the minh.
            else: 
                maxh.append(elem) #If the new element is greater than minh, add it to maxh.
        heapq.heapify(minh) #Heapify the list to maintain the properties.
        heapq._heapify_max(maxh)
    else:
        print ("List size is not valid.") #If the difference of size of the two lists is 2 or more.
    return maxh, minh #Return the minh and maxh.
                
#Test case:
ret = add_to_median_heap([7,8,9,10],[1,4,5,3],6)
minh = ret[1] #Adds the second value returned to the minh.
maxh = ret[0] #Adds the first value returned to the maxh
print (minh, maxh) #Print the two values.
```

Output:
```
[6, 7, 9, 10, 8] [5, 4, 1, 3]
```

## Question 2
Write a function `median(minh, maxh)`. It must return the median element.


<!-- grade_id: cell-43ce866dd1dc2d5b, points: 0, grade: false, solution: true -->
```python
def median(minh, maxh): #Function to find the median.
    if len(maxh) > len(minh): #If the length of maxheap is greater than minheap, return the value of maxheap
        return maxh[0]
    elif len(minh) > len(maxh):#If the length of minheap is greater than maxheap, return the value of minheap
        return minh[0]
    else:
        return (maxh[0] + minh[0])/2 #If both the lists have equal length, sum up the two prioritized elements and divide it by 2 to find the median.
median (minh, maxh)
```

Output:
```
6
```

<!-- grade_id: cell-673faa69bf330eda, points: 0, grade: true, solution: false -->
```python
# Please ignore this cell. This cell is for us to implement the tests 
# to see if your code works properly. 
```

## Question 3.

Uncomment and run the testing code given below to test your functions. It should print out numbers ranging from 1 to 50, in that order.

```python
minh = []
maxh = []
for a in range(1,100,2):
    add_to_median_heap(minh, maxh, a)
    print(float(median(minh, maxh)))  
```

Output:
```
1.0
2.0
3.0
4.0
5.0
6.0
7.0
8.0
9.0
10.0
11.0
12.0
13.0
14.0
15.0
16.0
17.0
18.0
19.0
20.0
21.0
22.0
23.0
24.0
25.0
26.0
27.0
28.0
29.0
30.0
31.0
32.0
33.0
34.0
35.0
36.0
37.0
38.0
39.0
40.0
41.0
42.0
43.0
44.0
45.0
46.0
47.0
48.0
49.0
50.0
```

## Question 4.
What’s the worst case complexity to build a median heap using `add_to_median_heap`?

<!-- grade_id: cell-c8e33ebdc8aaf328, points: 0, grade: true, solution: true -->
The worst case scenario for the heapify function is O(n). If we look at the code, we can see that at each run of the function, we need to ensure the correctness of the max and min heap properties. Rest of the operations have smaller order of growth (i.e. linear or logarithmic time), so we can consider the final complexity to be O(n).

## Question 5.
What’s the worst case complexity of `median`?

<!-- grade_id: cell-c061920faa337890, points: 0, grade: true, solution: true -->
There should be no worst case for the median, since it's runtime is not dependent on the size of the arrays. Therefore, we have a constant time factor, O(1).

## Question 6.

How does this way of finding the median compare with the vanilla way of sorting the list and pick the middle element? Use arguments based on efficiency or clearity of the respective algorithms.

<!-- grade_id: cell-1f31afecda9730a7, points: 0, grade: true, solution: true -->
The algorithms we used before to sort the elements, and find the median (i.e. merge sort, quicksort, etc.) had the best case and average case time complexity of O(n log(n)). And, as we showed before, the time complexity of this method is O(n). Therefore, this method is more efficient in terms of the order of growth.

## [Optional] Question 7.

Is it possible to extend this idea to any percentile? If it is, then write code to do so. If it’s not possible, prove why it is not possible.

<!-- grade_id: cell-4d0c197e6d484338, points: 0, grade: true, solution: true -->
```python
# # YOUR CODE HERE
# raise NotImplementedError()
```

## Part B. Quickselect

Quicksort can be modified to find the $k$-th smallest element in an unordered list. This is known as quickselect. It does this by choosing a partition (as in quicksort). Once the list has been partitioned then we know how many elements lie to the left and to the right of the partition. This allows us to recursively call quickselect on the correct sublist.

## Question 1.

Write a function `qselect(lst, k)`, which takes both a list and an index $k$. The function must then return the $k$-th smallest item in the list.


<!-- grade_id: cell-fb9c2bc9d5b4bc92, points: 0, grade: false, solution: true -->
```python
def qselect(lst,k):
        if k<0 or k>len(lst)-1:
            return print ("OUT OF INDEX")
        return selection(lst, k, 0, len(lst)-1)
        
def partition(A,p,r):
    pivot = A[r] #Takes last element of the array as the pivot
    index = p-1 #Creates the index as  -1.
    for i in range(p,r): #Iterated through the array
        if A[i] <= pivot: #Checks if the i-th element is samller than the pivot element
            index += 1 #Increments the index by 1 if true
            A[index],A[i] = A[i], A[index] #Swaps the two elements to ensure that the smaller elements are to the left and bigger ones to the right.
    A[index+1],A[r] = A[r], A[index+1] #Moves the pivot element after the smaller left part
    return index+1 #Returns the index of the pivot.

def selection(lst, k, l, h):
    if l == h:
        return arr[l]
    q = partition(lst, l, h)
    if k == q:
        return lst[k]
    elif k<q:
        return selection (lst, k, l, q-1)
    else:
        return selection(lst, k, q+1, h)
```

<!-- grade_id: cell-85a25ed6a1f62a66, points: 0, grade: true, solution: false -->
```python
import random
random.seed(123) # introducing a seed for reproducibility purposes
lst1 = list(range(100))
random.shuffle(lst1)
lst2 = []
for a in range(100):
    lst2.append(qselect(lst1, a))
assert(lst2 == list(range(100)))
```

## Question 2.
Uncomment and run the testing code given below to test your function. It should print out integers from 0 to 99.


```python
import random
random.seed(123) # introducing a seed for reproducibility purposes
lst = list(range(100))
random.shuffle(lst)
for a in range(100):
    print(qselect(lst, a))

```

Output:
```
0
1
2
3
4
5
6
7
8
9
10
11
12
13
14
15
16
17
18
19
20
21
22
23
24
25
26
27
28
29
30
31
32
33
34
35
36
37
38
39
40
41
42
43
44
45
46
47
48
49
50
51
52
53
54
55
56
57
58
59
60
61
62
63
64
65
66
67
68
69
70
71
72
73
74
75
76
77
78
79
80
81
82
83
84
85
86
87
88
89
90
91
92
93
94
95
96
97
98
99
```

## Question 3.

Write down the recurrence relation for your code.

<!-- grade_id: cell-24144607b3ffade3, points: 0, grade: true, solution: true -->
The recursion occurs in a subarray of length n/2, so it takes the form T(n/2). We also have the partition function whose time complexity is O(n). The function to find the median has constant time, so the big-O is O(1).

Finally, we can write the recursive relation as:
$T(n) = T(n/2) + O(n) + O(1)$
We can also write this as O(n).

## Question 4.

Solve the recurrence relation for quickselect in the best case.


<!-- grade_id: cell-c125f4742933e820, points: 0, grade: true, solution: true -->
The function can be expanded to understand the pattern.
$O(n) + O(n/2) + ... + O(1)$. This is equal to O(2n). We can safely disregard the constant, since we are only concerned with the order of growth. So, the complexity for the best case is O(n).

## Question 5.

Solve the recurrence relation for quickselect in the worst case.

<!-- grade_id: cell-852c3a71d8c2d7c1, points: 0, grade: true, solution: true -->
Worst case occurs when the partitioning is unbalanced, i.e. one sub array has n-1 elements, and the other has 0 elements. This happens when the pivot element we choose is either the largest or smallest value. The other comparison we need to make is to ensure the correctness of the $k^{th}$ element. We can write the relation as:
$T(n) = T(n-1) + O(n) \rightarrow O(n^2)$

```python

```