Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "Batsal Ghimire"
COLLABORATORS = ""
```

---

# CS110 Pre-class Work 1.2



## Question 1
The following pseudocode has been extracted from Cormen et al. You should analyze the lines of the pseudocode trying to understand what every line in the pseudocode does and how many times is that line executed. 
![Insertion sort pseudo-code](https://user-images.githubusercontent.com/44539656/72424387-d14ac780-377d-11ea-8e15-0b790bd06b29.png)

Answer the following questions: What does line 7 do? Why is it necessary?

<!-- grade_id: cell-3bac59ff24b11a27, points: 0, grade: true, solution: true -->
In the absence of this line, the while loop would never end and would run infinitely. This would not give us the required solution. So, in order to break from this loop after all the prior elements have been checked for their order, this line is used. One of the requirements for the while loop is that i>0, and if 'i' will be equal to 0, this loop will break. This step is however reversed back to its initial state, when in the for loop, the value of i is set to 'j-1'.

## Question 2
The following Python code is an attempt to implement the Insertion-Sort pseudocode above. However, this code has a bug. 


```python
def insertionSort(A):
    for j in range(len(A)):
        key = A[j]
        i= j-1
        while i >= 0 and A[i]>key:
            A[i+1] = A[i]
            i -= 1
            A[i+1] = key
    return A
```

Copy and paste the code in the cell below. Identify and correct the bug. 

<!-- grade_id: cell-1b7488e89cb501c5, points: 0, grade: false, solution: true -->
```python
def insertionSort(A):
    # YOUR CODE HERE
    for j in range(len(A)):
        key = A[j]
        i= j-1
        while i >= 0 and A[i]>key:
            A[i+1] = A[i]
            i -= 1
        A[i+1] = key
    return A
    raise NotImplementedError()

```

<!-- grade_id: cell-3432ce8ae2493fa6, points: 1, grade: true, solution: false -->
```python
assert(insertionSort([0]) == [0])
assert(insertionSort([-1,1]) == [-1,1])
assert(insertionSort([1,-1]) == [-1,1])
assert(insertionSort([1,6,3,6]) == [1,3,6,6])
```

How do you know this code works as intended now?

<!-- grade_id: cell-442998d5fdd8b561, points: 0, grade: true, solution: true -->
The best way to know that this code works is looking at the input and seeing if it produces the correct output. But, in this case, the output of both the buggy and the fixed code was the same. The only problem was with the indentation, so the value of 'key' was assigned to A[i+1] everytime the while loop ran.

## Question 3 
Implement the bubble sort algorithm in the cell below. First, look up the corresponding pseudocode in Cormen et al. (page 40), and write a Python implementation for the bubble sort.  Like in the `insertionSort`,  your new Python function, `bubbleSort`, takes a list of elements as an argument and returns the sorted list.

<!-- grade_id: cell-b709733c96d8b615, points: 0, grade: false, solution: true -->
```python
def bubbleSort(A):
    # YOUR CODE HERE
    for i in range (len(A)):
        for j in range(len(A)-1):
            if A[j]>A[j+1]:
                A[j+1],A[j] = A[j],A[j+1]
    return A
    raise NotImplementedError()
```

<!-- grade_id: cell-ffc9b251ea0047ba, points: 2, grade: true, solution: false -->
```python
# Please ignore this cell. This cell is for us to implement the tests 
# to see if your code works properly. 
```

Is your program bug-free? Think of three different [“test cases”](https://www.hiredintech.com/classrooms/algorithm-design/lesson/83) to check the correctness of your code. If a code passes the tests you have just outlined, is this an unambiguous proof that the code is correct? Give your answer in the cell below. Feel free to use additional cells (Markdown or Code) as you wish.

<!-- grade_id: cell-5ecb51404140cb2f, points: 0, grade: true, solution: true -->
```python
import  random
def rand_list(start, end, num): #This creates a random list with the given parameters in an unsorted order and sorts it.
    List = [] 
  
    for i in range(num): 
        List.append(random.randint(start, end)) 
  
    print ("Unsorted Random List:", List)
    print ("Sorted List",bubbleSort(List))
rand_list(10,40,10)

#The sorted list we obtain is always correct.

assert(bubbleSort([0]) == [0])
assert(bubbleSort([-1,1]) == [-1,1])
assert(bubbleSort([1,-1]) == [-1,1])
assert(bubbleSort([1,6,3,6]) == [1,3,6,6])

#We also don't get any assertion error when running these cases from the previous problem.
```

Output:
```
Unsorted Random List: [34, 27, 26, 20, 29, 29, 30, 36, 15, 26]
Sorted List [15, 20, 26, 26, 27, 29, 29, 30, 34, 36]
```

Although we have tested the code with a random list, this is not an unambigious proof that the code works for each and every case since there might be some edge cases or scenarios where the code might fail. But, this still works for majority of the test cases with integers.

## Question 4
The following pseudocode corresponds to a third sorting algorithm to be covered in Lesson 1.2, the selection sort.  Like in the previous step, write a Python implementation of the algorithm and make sure that your program is bug-free. `selectionSort` should take a list of elements as an argument and return the sorted list.

![Selection sort pseudo-code](https://user-images.githubusercontent.com/44539656/72424390-d14ac780-377d-11ea-9844-3fddb8e85794.png)

<!-- grade_id: cell-e04f94e5d418923c, points: 0, grade: false, solution: true -->
```python
def selectionSort(A):
    # YOUR CODE HERE
    n = len(A)
    for i in range (1,n-1):
        minidx = i
        for j in range (i+1,n):
            if A[j]<A[minidx]:
                minidx = j
        A[i],A[minidx] = A[minidx],A[i]
    return A
    raise NotImplementedError()
selectionSort([1,5,2,6,8,3])
```

Output:
```
[1, 2, 3, 5, 6, 8]
```

<!-- grade_id: cell-ffe912ac7922de70, points: 1, grade: true, solution: false -->
```python
# Please ignore this cell. This cell is for us to implement the tests 
# to see if your code works properly. 
```

**Please make sure you try your best to code insertion sort, selection sort and bubble sort, and bring these programs to class. You will need them for a class activity.**