Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "Batsal Ghimire"
COLLABORATORS = ""
```

---

# CS110 Pre-class Work 5.1

## Question 1.

### Question 1a.

Write the worst sort function in the world, `worst_sort`. This function takes a list, and then randomly shuffles it until it happens to be in sorted order. Once it is in sorted order then the list is returned.


<!-- grade_id: cell-ead9c74054a1c5da, points: 0, grade: false, solution: true -->
```python
import random    
def worst_sort(A):
    """
    Sort A in ascending order by randomly shuffling its 
    elements until they are in order.
    
    Input:
    - A: list of numerical values
    
    Output:
    - A: Sorted list
    """
    sorted_ = sorted(A) #Creates a sorted list of A for comparison
    while A != sorted_: #Checks if the list A is sorted after the random shuffle
        random.shuffle(A) #Shuffles the list randomly
    return A #Returns the sorted list A
```

<!-- grade_id: cell-ff4db3a3e8b04515, points: 1, grade: true, solution: false -->
```python
# Please ignore this cell. This cell is for us to implement the tests 
# to see if your code works properly. 
```

### Question 1b.
What is the best case complexity of this algorithm?

<!-- grade_id: cell-afac80171b3b6232, points: 0, grade: true, solution: true -->
The best case complexity would be O(1). This is possible if the first random shuffle we get would be sorted since we would not need to shuffle any further.

### Question 1c.

What is the average case complexity?


<!-- grade_id: cell-9d52d34daa1e3478, points: 0, grade: true, solution: true -->
We can use expected value to find the average case complexity but I'm  not sure how we can do it numerically. We could take the average number of times we found the right permutation while shuffling the list to find it. Also, there will be n! total permutation. So, we can consider O(n!) to be the worst case scenario.

### Question 1d.

For what size lists is this a feasible method?

<!-- grade_id: cell-ab40f08768579798, points: 0, grade: true, solution: true -->
This method is very inefficient, especially if we conside the size of the list to be larger. After 9 elements in the list, the runtime was no longer feasible and took very long time.

### [Optional] Question 1e.

Can you think of an even worse sorting method? In such a case, what would its complexity be? How big a list could you sort?

<!-- grade_id: cell-ac05806942341926, points: 0, grade: true, solution: true -->
YOUR ANSWER HERE

## Question 2.

Approximate median finder. Given a list and δ (a number between 0 and 0.5), the approximate median finder function returns a number that is guaranteed to lie between the (50-tilde{δ}/2)% and (50+\tilde{δ}/2)% percentiles, where \tilde{δ}=100*δ. Implement such a function by randomly selecting an element from the list and testing whether or not it lies within the bounds. If it doesn’t then retry with a new random element, but only a limited number of retries are allowed (you decide the maximum number of retries.)

### Question 2a.
Complete the following function


<!-- grade_id: cell-2ab65512a6148d3c, points: 0, grade: false, solution: true -->
```python
import numpy as np
def check_approx_med(A, num, delta):
    """
    Given a list and a number in a list, determine whether it is within (50-tilde{delta}/2)% and 
    (50+\tilde{delta}/2)% percentiles.
    X% percentile of a list is defined to be the smallest number in the list that is as large as at least
    X% numbers in the list. 
    
    Input:
    - A: list of numerical values
    - num: the number of interest
    - delta: the error, lies between 0 and 0.5
    
    Output:
    - isApproxMed: a boolean value indicating whether num is within the bound
    """
    n = len(A) #Assigns the length of A to n
    if n == 1: #Checks if the length of the list and return false if it is 1
        return False
    else: 
        lower = int((0.5-delta/2)*n) #Assigns the lower value as given by the question
        upper = int((0.5+delta/2)*n)#Assigns the upper value as given by the question
        for i in range (int(100//delta)): #Loops through the given number of times
            random_num = np.random.randint(lower, upper+1) #We randomly select values within the bounds
            if num == random_num: #Checks if we obtained the number of interest
                return True
```

<!-- grade_id: cell-df56d7433abb6517, points: 1, grade: true, solution: false -->
```python
assert(check_approx_med([0,1], 0, .25) == True)
assert(check_approx_med([0], 0, .25) == False)
```

### Question 2b.
Complete the following function that makes use of `check_approx_med` above.


<!-- grade_id: cell-e19685b545147dde, points: 0, grade: false, solution: true -->
```python
def approx_med_finder(A, delta):
    """
    Given a list, find a number in the list that is between 50-delta/2% and 50+delta/2% percentiles
    
    Input:
    - A: list of numerical values
    - delta: the error, lies between 0 and 0.5
    
    Output:
    - num: the approximated median, if it is found within 100//delta trials (this is one possible 
    maximum number of retries. While you are encouraged to play around with another limits, the 
    submitted work must use this number.)
    - None: finding failed, if nothing found within 100//delta trials

    Note: the 100//delta is chosen here as it is the expected number of trials for the first successful finding to occur (using geometric distribution)
    """
    for i in range (int(100//delta)):  #Iterates through the loop
        choice = np.random.randint(0, len(A)+1) #Generates a random number within the given bounds
        if check_approx_med(A, choice, delta): #Calls the previously created function with the given arguments
            return choice #Returns the choice
    return None
```

### Question 2c.

What is the probability of failure in each random trial? What is the probability of failure after all the allowed random trials? Does this scale with δ or N (the number of elements in a list)?


<!-- grade_id: cell-efe7ab76f6a2546a, points: 0, grade: true, solution: true -->
Delta sets the boundaries for the median, so the probability of each random trial is dependent on its value. As we increase the value of delta, the probability of failure decreases. 
The probability is also dependent on the number of allowed random trials. As we increase the number of trials, the probability should failure should decrease. If we say that the probability of one failure is (1-$\delta$). The probability of failure after x  random trials is $(1-\delta)^x$.

The number of elements 'N' will also affect the total probability of combined events whereas delta changes the probability of each event.

### Question 2d.
Analyze the expected runtime of `approx_med_finder`. Note that because the function uses `check_approx_med`, you will most likely need to analyze the runtime of that function, too.

<!-- grade_id: cell-86b199798477fc2a, points: 0, grade: true, solution: true -->
For approx_med_finder, we should have the right value after 1/$\delta$ trials. And, there is another function check_approx_med which run through all the elements, N. So, we can multiply these two values to get O(1/$\delta$)*(N) = O(N/$\delta$)

```python

```