Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "Batsal Ghimire"
COLLABORATORS = ""
```

---

# CS110 Pre-class Work 3.1

## Question 1. 
Paste in your Python implementation of the maximum subarray from the previous class in the cell below and use that to find out the value of the maximum subarray of this array: `A = [-2, -3, 4, -1, -2, 1]`

<!-- grade_id: cell-e69c8bbcd40ca242, points: 0, grade: true, solution: true -->
```python
#Code from the previous class.
import math
def max_sum_brute(A):
    n = len(A)
    sum_ = -math.inf
    for i in range (n+1):
        for j in range (i, n+1):
            if sum_<sum(A[i:j+1]):
                sum_ = sum(A[i:j+1])
                start = i
                end = j
    return sum_, start, end
    raise NotImplementedError()
arr = [-2, -3, 4, -1, -2, 1]
max_sum_brute(arr)
```

Output:
```
(4, 2, 2)
```

## Question 2. 
Now, your friend Joe comes and appends a single extra number at the end of the array, which becomes:  `B = [-2, -3, 4, -1, -2, 1, 8]`. Do you have to re-run the entire maximum subarray again? Explain your answer. 
The subsequent questions will help you figure out an efficient algorithmic strategy to address the last question, but make sure to write your explanation above first, before answering the remaining questions.


<!-- grade_id: cell-d4b7cd845c816263, points: 0, grade: true, solution: true -->
Yes, we need to re-run the entire maximum subarray. Using brute force, we are making comparisons between different sets of subarrays, so when we append an element at the end of the array, there are new subarrays that are formed. There might be new maximum subarray that has a higher sum than the original array.

```python
B = [-2, -3, 4, -1, -2, 1, 8]
max_sum_brute(B)
```

Output:
```
(10, 2, 6)
```

## Question 3. 

**Determine if the following statement is True or False and explain your answer.**
If the maximum subarray of the array A is different than the maximum subarray of the array B (questions 1 and 2), the new maximum-subarray doesn’t need to include 8 (i.e., the newly appended element). 


<!-- grade_id: cell-9377964a8756f13b, points: 0, grade: true, solution: true -->
This is false. If the new maximum-subarray would not include 8, then it would be the same as the initial array A and also should have the same maximum subarray. '8' is an appended element making it a part of the list, so it needs to included for us to obtain the correct results.

## Question 4.
Complete the Python function `incremental_max_subarray(x, mx)` in the cell below.

This [video](https://www.youtube.com/watch?v=AAgErqQmwmA&list=PLF_a-qBXTGFektoI6JUOTRL36JlvD04BR&index=4&t=0s)  might be helpful to understand the `incremental_max_subarray` problem.

<!-- grade_id: cell-0230e459f3d701f8, points: 0, grade: false, solution: true -->
```python
import math
def incremental_max_subarray(x, mx):
    """ 
    Inputs:
    - x: a NON-EMPTY list of numbers (e.g., the array B in the first two questions above). 
        * If x has 1 element: returns the value of the element regardless of the value of mx
    - mx: the maximum subarray of x excluding its last element (i.e., compute the 
    maximum subarray of the input array x considering only its first len(x) - 1 elements)
    
    Output: The maximum subarray of the array x.
    
    For example, using the array B in question 2, the result of incremental_max_subarray(B, 4) 
    is 10 (10 = 8 + 1 - 2 -1 + 4). 
    """
    n = len(x)  #Finds the length of the array.
    if n>1: #If the array has more than one element, the following lines will run.
        new_sum = 0 #Initialize the new sum as 0.
        for i in range (n-1, -1, -1): #Loop runs from 'n-1' to '-1' in the opposite direction i.e. decreases by -1.
            new_sum = new_sum + x[i] #Restate the value of new_sum by adding x[i]
            if new_sum>mx: #Compares if new_sum is greater than the initial maximum sum.
                mx = new_sum #If true, then we will have a new maximum sum.
        return (mx) #Returns the new maximum.
        raise NotImplementedError()
        
    else:
        return x[0] #If there is only one element, it has to be the maximum.
```

<!-- grade_id: cell-9abf19e56620ffa3, points: 1, grade: true, solution: false -->
```python
B = [-2, -3, 4, -1, -2, 1, 8]
assert(incremental_max_subarray(B, 4) == 10)
assert(incremental_max_subarray(B[:1], 0) == B[0])
```

<!-- grade_id: cell-bacebe71191cbd0f, points: 0, grade: true, solution: true -->
YOUR ANSWER HERE

<!-- grade_id: cell-fd96d4ccccd99fe6, points: 0, grade: false, solution: true -->
```python
def max_subarray(A):
    """
    Using `incremental_max_subarray` iteratively on A to produce the value of the maximum
    subarray of A.
    
    Inputs:
    - A: a NON-EMPTY list of floats
    
    Outputs: flaot, the sum of the maximum subarray of A
    """
    # YOUR CODE HERE
    n = len(A) #Find the length of the array.
    if n == 1: #Runs if there is only one element in the array. 
        return float(A[0]) #Returns the only element as the maximum.
    else:
        mx = A[0] #Sets mx to be the value of the maximum.
        for i in range (1,n):
            mx = incremental_max_subarray(A[:i+1],mx) #Calls the previous function with the current max, and the array with the added element.
    return float(mx) #Returns the maximum as a float.
```

<!-- grade_id: cell-669ad779766aa2c3, points: 1, grade: true, solution: false -->
```python
# Please ignore this cell. This cell is for us to implement the tests 
# to see if your code works properly. 
```

### Part 2.
Is this more efficient than the divide-and-conquer approach? Explain.

<!-- grade_id: cell-cd8f0b130a7136db, points: 0, grade: true, solution: true -->
The efficiency compared with divide and conquer would depend on the information that is provided. Like in question 4, if we are provided with the initial maximum and we have to calculate the new maximum with the addition of a new element at the end of the list, this method would be more efficient. Since, we have to go through each elements only once, we get the efficiency to be $\Theta (n)$, whereas the efficiency of the divide and conquer is $\Theta (n lg(n))$.

But, if we need to find the maximum sum from the entire array without any information on the initial maximum, the efficiency would still be $\Theta (n^2)$. In this case, the divide and conquer would be more efficient.

```python

```