Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "Batsal Ghimire"
COLLABORATORS = ""
```

---

# CS110 Pre-class Work 3.2

## Question 1.
Given the array `H=[39, 85, 85, 16, 49, 7, 49, 92, 76, 15, 21, 30, 29, 31, 28]`, perform the following operations:
1. Draw the corresponding binary tree of H. Is the binary tree a valid max heap? Explain your answer.
2. Using as a model the drawing examples illustrated in Figure 6.2 of Cormen et al.,  draw a step-by-step transformation of the array above into a valid max heap. 
3. Now that you have obtained a valid max heap, write out the corresponding array that stores the valid max-heap.

Use as many cells as you wish for this question.

<!-- grade_id: cell-b1f821b034b44619, points: 0, grade: true, solution: true -->
1. The binary tree is not a valid max heap because some of the parent nodes are smaller that their respective children. And, for it to be a valid max heap, the parent node should always be greater than their child node.

2. ![1](https://user-images.githubusercontent.com/44539656/73351181-55c13e00-4286-11ea-8cca-7ad613989ab6.jpg)
    ![2](https://user-images.githubusercontent.com/44539656/73351182-55c13e00-4286-11ea-86ce-76b6ce9a42db.jpg)
    ![3](https://user-images.githubusercontent.com/44539656/73351184-55c13e00-4286-11ea-8803-1cdb44739e1d.jpg)
    ![4](https://user-images.githubusercontent.com/44539656/73351185-5659d480-4286-11ea-8660-b073ee90a2be.jpg)
    ![5](https://user-images.githubusercontent.com/44539656/73351186-5659d480-4286-11ea-879c-88feafe0e16b.jpg)
    ![6](https://user-images.githubusercontent.com/44539656/73351188-56f26b00-4286-11ea-97fc-e2e7b738f3e2.jpg)
    ![7](https://user-images.githubusercontent.com/44539656/73351198-5c4fb580-4286-11ea-8caf-9bd934e360f0.jpg)

```python
#3. 
new_array = [92,85,85,76,39,30,49,16,49,15,21,7,29,31,28]
```

## Question 2. 
Consider the following questions on the $MAX-HEAPIFY$ operation.
### Question 2a.

In the pseudocode of $MAX-HEAPIFY$ (Cormen et al., p.154, or you can view it [here](https://drive.google.com/open?id=1e_3jsX4-qQCfZXKMok_T6LPFh9FwtmT5)), what does A.heap-size mean and what is the idea behind the local variable largest? 


<!-- grade_id: cell-06106b81a909b9e6, points: 0, grade: true, solution: true -->
A.heap.size gives us the number of elements included in the heap. If we are given an arrary, and we take a certain number to form the binary tree, it is reprsented by the given term. It should always be less than or equal to the length of the arra, A.

The local variable largest helps us to ensure that the parent node is always greater than the child node. If holds the value of the largest index, so we can swap the two elements.

### Question 2b.
The functions $LEFT(i)$ and $RIGHT(i)$, lines 1 and 2 in the $MAX-HEAPIFY$ pseudocode, return the array index of the left and right child, respectively, of a node in a binary tree. From reading Section 6.1, you know that the input to both functions is an integer number, $i$, which corresponds to the array index of the parent node in the array. Review Section 6.1 for more information. Write a Python implementation of the functions $LEFT(i)$ and $RIGHT(i)$ by filling in the cells below.

<!-- grade_id: cell-2efd4017321c8d48, points: 0, grade: false, solution: true -->
```python
def left(i):
    # YOUR CODE HERE
    return 2*i+1
    raise NotImplementedError()
```

<!-- grade_id: cell-507ddbab62e6a32c, points: 1, grade: true, solution: false -->
```python
# Please ignore this cell. This cell is for us to implement the tests 
# to see if your code works properly. 
```

<!-- grade_id: cell-b7e5ceeacaeaf1b8, points: 0, grade: false, solution: true -->
```python
def right(i):
    # YOUR CODE HERE
    return 2*1+2
    raise NotImplementedError()
```

<!-- grade_id: cell-c24867515653788a, points: 1, grade: true, solution: false -->
```python
# Please ignore this cell. This cell is for us to implement the tests 
# to see if your code works properly. 
```

### Question 2c.
Write a Python implementation of the MAX-HEAPIFY operation using the pseudocode above, and your newly written functions, `left` and `right`.

<!-- grade_id: cell-52289253a243341b, points: 0, grade: false, solution: true -->
```python
def heapify(heap, i):
    """
    Inputs:
    - heap: a list of floats. Assume that the heap size is the length of the heap.
    
    No output is needed. This function should modify (if necessary) heap in-place.
    """
    
    # YOUR CODE HERE
    n = len(heap) #Finds out the length of the heap.
    l = left(i) #Calls on the above function to define the left child.
    r = right(i) #Calls on the above function to defind the right child.
    largest = i #Sets the largest value to i.
    if l <n and heap[l]>heap[i]: #Checks if the left child is greater than the parent.
        largest = l #If true changes the largest variable to l.
    else:
        largest = i
        
    if heap[r]>heap[largest] and r<n: #Checks if the right child is greater than the largest value.
        largest = r
        
    if largest != i: #Checks if the parent is the largest element
        heap[i],heap[largest] = heap[largest],heap[i] #Swaps the values if true
        heapify(heap, largest) #Calls the function recursively until the condition fails.
```

<!-- grade_id: cell-7f63745ff34b881b, points: 1, grade: true, solution: false -->
```python
A = [39, 85, 85, 16, 49, 7, 49, 92, 76, 15, 21, 30, 29, 31, 28]
heapify(A,0)
assert(A == [85, 49, 85, 16, 39, 7, 49, 92, 76, 15, 21, 30, 29, 31, 28])

```

## Question 3. 
Next, write a Python implementation of the BUILD_MAX_HEAP operation using the pseudocode provided in Section 6.3 of Cormen et. al. Test your Python implementation using the array in problem 1, and make sure your Python codes produce a valid max heap.

<!-- grade_id: cell-9237da6268a5e3eb, points: 0, grade: false, solution: true -->
```python
def build_max_heap(A):
    """
    Input:
    - A: a list of floats.
    
    No output is needed. The function should turn A into a valid max heap, in-place.
    """
    # YOUR CODE HERE
    heap_size = len(A) #Sets the length of the array A to heap_size
    for i in range ((heap_size//2), -1, -1): #Iteration from the mid index to the first index
        heapify(A, i)
A = [4, 1, 3, 2, 16]
print (build_max_heap(A))
#It gives me a maximum recursion depth exceeded, but it works with shorter arrays.
```

Output:
```
None
```

<!-- grade_id: cell-5dc147b837da3f7e, points: 1, grade: true, solution: false -->
```python
A = [4, 1, 3, 2, 16, 9, 10, 14, 8, 7]
build_max_heap(A)
assert(A == [16, 14, 10, 8, 7, 9, 3, 2, 4, 1])
```

Output:
```
... last 1 frames repeated, from the frame below ...

<ipython-input-7-f00d87b991f2> in heapify(heap, i)
     22     if largest != i: #Checks if the parent is the largest element
     23         heap[i],heap[largest] = heap[largest],heap[i] #Swaps the values if true
---> 24         heapify(heap, largest) #Calls the function recursively until the condition fails.

RecursionError: maximum recursion depth exceeded while calling a Python object
```

## Question 4. 

Lastly, write Python implementations of the $MIN-HEAPIFY$ and $BUILD-MIN-HEAP$ operations for a min heap data structure. You can use your $MAX-HEAPIFY$ and $BUILD-MAX-HEAP$ Python function as models, just remember that the latter two functions support operations for the max heap data structure. Test your Python implementation of $BUILD-MIN-HEAP$ using the array in problem 1, and make sure your Python codes produce a valid min heap. 

<!-- grade_id: cell-d86943f127876fc8, points: 0, grade: false, solution: true -->
```python
def min_heapify(heap, i):
    """
    Inputs:
    - heap: a list of floats. Assume that the heap size is the length of the heap.
    
    No output is needed. This function should modify (if necessary) heap in-place.
    """
    # YOUR CODE HERE
    n = len(heap) 
    l = left(i) 
    r = right(i)
    smallest = i
    
    if heap[l]<heap[i] and l<n:
        smallest = l 
    else:
        smallest = i
    if heap[r]<heap[smallest] and r<n:
        smallest = r
        
    if smallest != i: 
        heap[i],heap[smallest] = heap[smallest],heap[i]
        heapify(heap, smallest) 
```

<!-- grade_id: cell-6efa45e3956ad0b8, points: 0, grade: false, solution: true -->
```python
def build_min_heap(A):
    """
    Input:
    - A: a list of floats.
    
    No output is needed. The function should turn A into a valid min heap, in-place.
    """
    # YOUR CODE HERE
    heap_size = len(A) #Sets the length of the array A to heap_size
    for i in range (heap_size//2, -1, -1): #Iteration from the mid index to the first index
        min_heapify(A, i)
    
```

<!-- grade_id: cell-7e64941402f819c2, points: 1, grade: true, solution: false -->
```python
# Please ignore this cell. This cell is for us to implement the tests 
# to see if your code works properly. 
```

```python

```

```python

```

```python

```